# Assignment 1
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [datasets/NISPUF17.csv](datasets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [datasets/NIS-PUF17-DUG.pdf](datasets/NIS-PUF17-DUG.pdf).

## Question 1 (20%)
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-d4625c53be146b54, points: 0, grade: false, solution: true -->
```python
#https://www.cdc.gov/vaccines/imz-managers/nis/downloads/NIS-PUF17-DUG.pdf

import pandas as pd
import numpy as np

def proportion_of_education():
    df = pd.read_csv('datasets/NISPUF17.csv')
    len_df = len(df)
    edu = df["EDUC1"]
    edu1 = edu==1
    edu2 = edu==2
    edu3 = edu==3
    edu4 = edu==4
    edu_dict = {"less than high school":len(edu[edu1])/len_df, "high school":len(edu[edu2])/len_df, "more than high school but not college":len(edu[edu3])/len_df, "college":len(edu[edu4])/len_df}
    return edu_dict
    
    
proportion_of_education()
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-75b071a5208d6122, points: 0, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
```

## Question 2 (20%)

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-2cbda7565b40aaf3, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

def average_influenza_doses():
    #read csv file
    df = pd.read_csv('datasets/NISPUF17.csv')
    
    #get the relevent columns
    new_data= pd.DataFrame(df, columns=['CBF_01', 'P_NUMFLU'])
    
    #create to groups of breastfed and not breastfed
    breastfed = new_data.loc[new_data['CBF_01'] == 1]
    not_breastfed = new_data.loc[new_data['CBF_01'] == 2]
    
    #use the breastfed only and find average num vaccines
    vacB = breastfed['P_NUMFLU'].dropna()
    sum_vacB = sum(vacB)
    count_vacB = len(vacB)
    avg_vacB = sum_vacB/count_vacB
    
    #use the NOT breastfed only and find average num vaccines
    vacNb = not_breastfed['P_NUMFLU'].dropna()
    sum_vacNb = sum(vacNb)
    count_vacNb = len(vacNb)
    avg_vacNb = sum_vacNb/count_vacNb
    
    #tuple averages
    averages = (avg_vacB,avg_vacNb)
    
    #return averages
    return averages

    '''
  
    for group, frame in df.where(df["review_scores_location"=='t']>9).dropna(how='all').groupby("review_scores_value")
    '''    

average_influenza_doses()
```

<!-- grade_id: cell-744f239f462ce2c9, points: 20, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."
```

## Question 3 (25%)
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value I'm looking for starts with the digits `0.00773`.

<!-- grade_id: cell-28bc178306233102, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

def chickenpox_by_sex():
    #read csv file
    df = pd.read_csv('datasets/NISPUF17.csv')
    
    #get the relevent columns
    new_data= pd.DataFrame(df, columns=['SEX', 'HAD_CPOX', 'P_NUMVRC'])
    
    #create to groups of male and female
    male = new_data.loc[new_data['SEX'] == 1]
    female = new_data.loc[new_data['SEX'] == 2]
    
    
    #-------------------------MALE-------------------------
    #use male only and vaciniated
    vacM0 = (male['P_NUMVRC']>=1).dropna()
    vacM1= ((male[vacM0]!=77) & (male[vacM0]!=99)).dropna()
    vacM2 = (male[vacM1]).dropna()
    
    vacMYes = vacM2
    
    
    #-----male had pox-----
    poxM = vacMYes['HAD_CPOX']==1
    poxMYes = vacMYes[poxM]
    
    #number males who were vaccinated and contracted CPOX
    vacMPoxYes = len(poxMYes)
    
    #----male DID NOT HAVE pox-----
    poxMx = vacMYes['HAD_CPOX']==2
    poxMNo = vacMYes[poxMx]
    
    #number males who were vaccinated and DID NOT contracted CPOX
    vacMPoxNo = len(poxMNo)
    
    
   #***RATIO MALE number of males who contracted chickenpox but were vaccinated against it VS those who were vaccinated but did not contract chicken pox 
    ratioMale = vacMPoxYes/vacMPoxNo
    
    
    #-------------------------FEMALE-------------------------
     #use female only and vaciniated
    vacF0 = female['P_NUMVRC']>=1
    vacF1= ((female[vacF0]!=77) & (female[vacF0]!=99)).dropna()
    vacF2 = (female[vacF1]).dropna()
    
    vacFYes = vacF2
    
    
    #-----female had pox-----
    poxF = vacFYes['HAD_CPOX']==1
    poxFYes = vacFYes[poxF]
    
    #number females who were vaccinated and contracted CPOX
    vacFPoxYes = len(poxFYes)
    
    #----female DID NOT HAVE pox-----
    poxFx = vacFYes['HAD_CPOX']==2
    poxFNo = vacFYes[poxFx]
    
    #number females who were vaccinated and DID NOT contracted CPOX
    vacFPoxNo = len(poxFNo)
    
    
   #***RATIO FEMALE number of females who contracted chickenpox but were vaccinated against it VS those who were vaccinated but did not contract chicken pox 
    ratioFemale = vacFPoxYes/vacFPoxNo
    
    
    #----------!!!!!!RETURN DICT!!!!!!----------
    vacPoxBySex = {'male':ratioMale, 'female':ratioFemale}
    return vacPoxBySex


chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-fba9cff0bf962f1a, points: 25, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."
```

## Question 4 (25%)
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8a03b576e9dc513f, points: 0, grade: false, solution: true -->
```python
#collaborated with Emily Kanterman
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    #read csv
    df = pd.read_csv('datasets/NISPUF17.csv')
    
    #smaller df with only colomns had pox or not and number of vaccines, drop rows that have NaN in these columns
    new_data= pd.DataFrame(df, columns=['HAD_CPOX', 'P_NUMVRC']).dropna()
    
    
    clean_df = ((clean_df2['HAD_CPOX']>=1) &(clean_df2['HAD_CPOX']<3))
    clean_df1 = new_data[clean_df]
    
                           
    corr, pval=stats.pearsonr(clean_df1["HAD_CPOX"], clean_df1['P_NUMVRC'])
    
   
    return corr, pval
    
corr_chickenpox()


```

Output:
```
(0.07044873460147985, 2.778026318363451e-18)
```

<!-- grade_id: cell-6e19cbdf20cb9e56, points: 25, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."
```

## Question 5 (10%)

Is there a relationship between states which voted for the republican candidate for president in 2016 (Trump) versus those who voted for the democratic candidate for president in 2016 (Hillary) and the rate of vaccination in those states in 2017? I would like you to answer this by doing a correlation test between the ratio of votes for Trump/Clinton against the ratio of Vaccination/Non-Vaccination (of Chickenpox!) by state.

Note: You are certainly using joins or merges here.

Note: You will might need the state [FIPS identifiers](https://en.wikipedia.org/wiki/Federal_Information_Processing_Standard_state_code#FIPS_state_codes) and will certainly want the [MIT Elections Data](https://dataverse.harvard.edu/dataset.xhtml?persistentId=doi:10.7910/DVN/42MVDX) which has the results from 2018.

Note: I don't know the answer to this question already.

<!-- grade_id: cell-7bfd291ce527a051, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
def corr_election():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    df_vac = pd.read_csv('datasets/NISPUF17.csv')
    df_vot = pd.read_csv('datasets/1976-2016-president.csv')
   
    #only 2016
    df_year = df_vot['year']==2016
    apply_year = df_vot[df_year]
    
    #only Trump and Hillary
    df_can = (apply_year['candidate']=='Trump, Donald J.') | (apply_year['candidate']=='Clinton, Hillary')
    apply_can = apply_year[df_can]
    
    #new smaller dataframe
    apply_can = pd.DataFrame(apply_can, columns=['state','state_fips','candidate','candidatevotes'])
    apply_can = apply_can.dropna()
    
    #new smaller dataframe for vaccine data
    vacc_df= pd.DataFrame(df_vac, columns=['STATE', 'P_NUMVRC'])
    vacc_df = vacc_df.dropna()
    
    #rename column in order to join on states
    election_df = apply_can.rename(columns={'state_fips':'STATE'})
    
    election_df.groupby(['STATE','state','candidate']).sum()

    voted_trump = election_df['candidate']=='Trump, Donald J.'
    t = election_df[voted_trump]
    trump_votes = t['candidatevotes']

    voted_clinton = election_df['candidate']== 'Clinton, Hillary'
    c = election_df[voted_clinton]
    clinton_votes = c['candidatevotes']            
   
    election_df['trump votes'] = trump_votes
    election_df['clinton votes'] = clinton_votes
    w = election_df.groupby(['state','STATE'])
    election_df2 = w.sum()

    election_df2['ratio trump/clinton'] = election_df2['trump votes']/election_df2['clinton votes']

    #new smaller dataframe for vaccine data
    vacc_df= pd.DataFrame(df_vac, columns=['STATE', 'P_NUMVRC'])
    vacc_df = vacc_df.dropna()
    
    vacc_yes = vacc_df['P_NUMVRC'] >= 1
    y_df = vacc_df[vacc_yes]
    y = y_df.groupby(['STATE'])
    yes = y.count()

    vacc_no = vacc_df['P_NUMVRC'] == 0
    n_df = vacc_df[vacc_no]
    n = n_df.groupby(['STATE'])
    no = n.count()

    merge_vacc = pd.merge(yes, no, how='outer', on='STATE')
    
    merge_vacc['ratio vacc/not-vacc']= merge_vacc['P_NUMVRC_x']/merge_vacc['P_NUMVRC_y']
    
    #merge the two data frames together
    merge = pd.merge(election_df2, merge_vacc, how='inner', on='STATE')

    corr, pval=stats.pearsonr(merge['ratio trump/clinton'],merge['ratio vacc/not-vacc'])
    
    #return corr, pval
    return corr, pval

corr_election()


```

Output:
```
(-0.2801083542361783, 0.046499248838739134)
```

<!-- grade_id: cell-2451528487f5b391, points: 10, grade: true, solution: false -->
```python

```

```python
import scipy.stats as stats
import numpy as np
import pandas as pd
```

```python
df_vac = pd.read_csv('datasets/NISPUF17.csv')
df_vot = pd.read_csv('datasets/1976-2016-president.csv')
```

```python
#only 2016
df_year = df_vot['year']==2016
apply_year = df_vot[df_year]
apply_year
```

Output:
```
year          state state_po  state_fips  state_cen  state_ic  \
3395  2016        Alabama       AL           1         63        41   
3396  2016        Alabama       AL           1         63        41   
3397  2016        Alabama       AL           1         63        41   
3398  2016        Alabama       AL           1         63        41   
3399  2016        Alabama       AL           1         63        41   
3400  2016         Alaska       AK           2         94        81   
3401  2016         Alaska       AK           2         94        81   
3402  2016         Alaska       AK           2         94        81   
3403  2016         Alaska       AK           2         94        81   
3404  2016         Alaska       AK           2         94        81   
3405  2016         Alaska       AK           2         94        81   
3406  2016         Alaska       AK           2         94        81   
3407  2016        Arizona       AZ           4         86        61   
3408  2016  
```

```python
#only Trump and Hillary
df_can = (apply_year['candidate']=='Trump, Donald J.') | (apply_year['candidate']=='Clinton, Hillary')
apply_can = apply_year[df_can]
apply_can
```

Output:
```
year                 state state_po  state_fips  state_cen  state_ic  \
3395  2016               Alabama       AL           1         63        41   
3396  2016               Alabama       AL           1         63        41   
3400  2016                Alaska       AK           2         94        81   
3401  2016                Alaska       AK           2         94        81   
3407  2016               Arizona       AZ           4         86        61   
3408  2016               Arizona       AZ           4         86        61   
3417  2016              Arkansas       AR           5         71        42   
3418  2016              Arkansas       AR           5         71        42   
3425  2016            California       CA           6         93        71   
3426  2016            California       CA           6         93        71   
3431  2016              Colorado       CO           8         84        62   
3432  2016              Colorado       CO           8         84      
```

```python
#new smaller dataframe
apply_can = pd.DataFrame(apply_can, columns=['state','state_fips','candidate','candidatevotes'])
apply_can = apply_can.dropna()
```

```python
apply_can
```

Output:
```
state  state_fips         candidate  candidatevotes
3395               Alabama           1  Trump, Donald J.         1318255
3396               Alabama           1  Clinton, Hillary          729547
3400                Alaska           2  Trump, Donald J.          163387
3401                Alaska           2  Clinton, Hillary          116454
3407               Arizona           4  Trump, Donald J.         1252401
3408               Arizona           4  Clinton, Hillary         1161167
3417              Arkansas           5  Trump, Donald J.          684872
3418              Arkansas           5  Clinton, Hillary          380494
3425            California           6  Clinton, Hillary         8753788
3426            California           6  Trump, Donald J.         4483810
3431              Colorado           8  Clinton, Hillary         1338870
3432              Colorado           8  Trump, Donald J.         1202484
3451           Connecticut           9  Clinton, Hillary          897572
```

```python
#rename column in order to join on states
election_df = apply_can.rename(columns={'state_fips':'STATE'})
```

```python
election_df.groupby(['STATE','state','candidate']).sum()

```

Output:
```
candidatevotes
STATE state                candidate                       
1     Alabama              Clinton, Hillary          729547
                           Trump, Donald J.         1318255
2     Alaska               Clinton, Hillary          116454
                           Trump, Donald J.          163387
4     Arizona              Clinton, Hillary         1161167
                           Trump, Donald J.         1252401
5     Arkansas             Clinton, Hillary          380494
                           Trump, Donald J.          684872
6     California           Clinton, Hillary         8753788
                           Trump, Donald J.         4483810
8     Colorado             Clinton, Hillary         1338870
                           Trump, Donald J.         1202484
9     Connecticut          Clinton, Hillary          897572
                           Trump, Donald J.          673215
10    Delaware             Clinton, Hillary          235603
                         
```

```python
voted_trump = election_df['candidate']=='Trump, Donald J.'
t = election_df[voted_trump]
trump_votes = t['candidatevotes']
#print(trump_votes)

voted_clinton = election_df['candidate']== 'Clinton, Hillary'
c = election_df[voted_clinton]
clinton_votes = c['candidatevotes']
```

```python
election_df['trump votes'] = trump_votes
election_df['clinton votes'] = clinton_votes
w = election_df.groupby(['state','STATE'])
election_df2 = w.sum()

election_df2


```

Output:
```
candidatevotes  trump votes  clinton votes
state                STATE                                            
Alabama              1             2047802    1318255.0       729547.0
Alaska               2              279841     163387.0       116454.0
Arizona              4             2413568    1252401.0      1161167.0
Arkansas             5             1065366     684872.0       380494.0
California           6            13237598    4483810.0      8753788.0
Colorado             8             2541354    1202484.0      1338870.0
Connecticut          9             1570787     673215.0       897572.0
Delaware             10             420730     185127.0       235603.0
District of Columbia 11             295553      12723.0       282830.0
Florida              12            9122861    4617886.0      4504975.0
Georgia              13            3967067    2089104.0      1877963.0
Hawaii               15             395738     128847.0       266891.0
Idaho                16           
```

```python
election_df2['ratio trump/clinton'] = election_df2['trump votes']/election_df2['clinton votes']
election_df2
```

Output:
```
candidatevotes  trump votes  clinton votes  \
state                STATE                                               
Alabama              1             2047802    1318255.0       729547.0   
Alaska               2              279841     163387.0       116454.0   
Arizona              4             2413568    1252401.0      1161167.0   
Arkansas             5             1065366     684872.0       380494.0   
California           6            13237598    4483810.0      8753788.0   
Colorado             8             2541354    1202484.0      1338870.0   
Connecticut          9             1570787     673215.0       897572.0   
Delaware             10             420730     185127.0       235603.0   
District of Columbia 11             295553      12723.0       282830.0   
Florida              12            9122861    4617886.0      4504975.0   
Georgia              13            3967067    2089104.0      1877963.0   
Hawaii               15             395738     128847.0       2668
```

```python
#new smaller dataframe for vaccine data
vacc_df= pd.DataFrame(df_vac, columns=['STATE', 'P_NUMVRC'])
vacc_df = vacc_df.dropna()
vacc_df
```

Output:
```
STATE  P_NUMVRC
3         27       1.0
4         34       0.0
6         34       1.0
7          9       0.0
8         42       0.0
10        20       1.0
12        48       1.0
16        20       2.0
19        31       1.0
20         4       1.0
21        29       1.0
22        39       1.0
24        32       1.0
25        56       1.0
26        36       1.0
27        25       1.0
28        47       1.0
29        48       1.0
30        27       1.0
34        25       1.0
35        38       1.0
36        23       0.0
39        27       1.0
40        17       1.0
41        23       1.0
43        46       1.0
44         6       1.0
46         9       1.0
47         9       1.0
49        27       1.0
...      ...       ...
28397     44       1.0
28398     30       0.0
28400     48       1.0
28402     49       1.0
28403     48       1.0
28404      1       1.0
28406     35       1.0
28407     22       1.0
28408     30       0.0
28409     22       1.0
28410     24       1.0
28411     35      
```

```python

vacc_yes = vacc_df['P_NUMVRC'] >= 1
#vacc_yes
y_df = vacc_df[vacc_yes]
#y_df
y = y_df.groupby(['STATE'])
yes = y.count()
#yes


vacc_no = vacc_df['P_NUMVRC'] == 0
n_df = vacc_df[vacc_no]
n = n_df.groupby(['STATE'])

no = n.count()
no

```

Output:
```
P_NUMVRC
STATE          
1            22
2            39
4            23
5            24
6            22
8            32
9            13
10           16
11           12
12           32
13           22
15           26
16           20
17           34
18           36
19           18
20           24
21           15
22           26
23           30
24           21
25           10
26           15
27           19
28           20
29           30
30           19
31           13
32           18
33           18
34           33
35           22
36           52
37           25
38           14
39           36
40           24
41           26
42           46
44           15
45           24
46           24
47           13
48          124
49           35
50           29
51           15
53           27
54           28
55           25
56           32
```

```python
merge_vacc = pd.merge(yes, no, how='outer', on='STATE')
merge_vacc
```

Output:
```
P_NUMVRC_x  P_NUMVRC_y
STATE                        
1             273          22
2             212          39
4             242          23
5             234          24
6             302          22
8             232          32
9             253          13
10            216          16
11            257          12
12            282          32
13            277          22
15            210          26
16            178          20
17            439          34
18            220          36
19            208          18
20            168          24
21            214          15
22            256          26
23            241          30
24            340          21
25            261          10
26            222          15
27            230          19
28            227          20
29            199          30
30            197          19
31            233          13
32            217          18
33            218          18
34            323          33
35            211
```

```python
merge_vacc['ratio vacc/not-vacc']= merge_vacc['P_NUMVRC_x']/merge_vacc['P_NUMVRC_y']
merge_vacc
```

Output:
```
P_NUMVRC_x  P_NUMVRC_y  ratio vacc/not-vacc
STATE                                             
1             273          22            12.409091
2             212          39             5.435897
4             242          23            10.521739
5             234          24             9.750000
6             302          22            13.727273
8             232          32             7.250000
9             253          13            19.461538
10            216          16            13.500000
11            257          12            21.416667
12            282          32             8.812500
13            277          22            12.590909
15            210          26             8.076923
16            178          20             8.900000
17            439          34            12.911765
18            220          36             6.111111
19            208          18            11.555556
20            168          24             7.000000
21            214          15         
```

```python
#merge the two data frames together
merge = pd.merge(election_df2, merge_vacc, how='inner', on='STATE')

    
```

```python
merge
```

Output:
```
candidatevotes  trump votes  clinton votes  ratio trump/clinton  \
STATE                                                                    
1             2047802    1318255.0       729547.0             1.806950   
2              279841     163387.0       116454.0             1.403018   
4             2413568    1252401.0      1161167.0             1.078571   
5             1065366     684872.0       380494.0             1.799955   
6            13237598    4483810.0      8753788.0             0.512214   
8             2541354    1202484.0      1338870.0             0.898134   
9             1570787     673215.0       897572.0             0.750040   
10             420730     185127.0       235603.0             0.785758   
11             295553      12723.0       282830.0             0.044985   
12            9122861    4617886.0      4504975.0             1.025064   
13            3967067    2089104.0      1877963.0             1.112431   
15             395738     128847.0       2668
```

```python
corr, pval=stats.pearsonr(merge['ratio trump/clinton'],merge['ratio vacc/not-vacc'])
corr, pval
```